# Homework 3: More String, Lists, and Python Objects (18 pts)

name: Jiayi Guo

email: jiayiguo@umich.edu

This homework assignment took me 10 hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I did not discuss this homework with anyone.

---

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. 


Before submitting, please make sure to __Cells->Run All__ executes without errors; errors in your code translate directly to point deductions. 
In general, you don't need to do explicitly raise errors (e.g. with the ```raise``` function) if we don't ask you to in the problem statement.
However, even in cases where we ask you to check for errors, your submission should not contain any examples of your functions actually raising those errors.

Note that many parts of this homework where you are expected to type in code will have ```NotImplementedError()``` as a placeholder. You need to delete this function and replace it with your own code.

## Importing libaries

Unless specifically prohibited, you can import any Python library that comes included with Python (e.g. ```math```, ```datetime```) or that has been used in the lecture or homework prompts. You do not need any other libraries to complete the homework. Note that the graders will not install any additional libraries to grade an assignment.


<!-- ## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking: -->

## Homework tips 

1. **Start early!** If you run into trouble installing things or importing packages, it’s
best to find those problems well in advance, not the night before your assignment is
due when we cannot help you!

2. **Make sure you back up your work!** At a minimum, do your work in a Dropbox
folder. Better yet, use git, which is well worth your time and effort to learn.

3. **Be careful to follow directions!** Remember that Python is case sensitive. If
you are ask you to define a function called my_function and you define a function
called My_Function, you will not receive full credit. You may want to copy-paste
the function names below to make sure that the functions in your notebook match.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:


## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. You will notice some `read-only` cells in the assignment that contain `assert` statements. These are tests that your code must pass for your solution to be correct. If any of the tests fail, you will get an python error and not get points for that question. 

**Note:** The tests shown not are not comprehensive; additional tests will be used at grading time. You are encouraged to read the problem carefully and verify your code covers all possible cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

```python
%load_ext pycodestyle_magic
```

```python
%%pycodestyle
import collections
import string
import itertools
import math
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

<!-- grade_id: primes, points: 0, grade: false, solution: true -->
```python
## %%pycodestyle
def primes():
    # YOUR CODE HERE
    def is_prime(n):
        # n is prime number(>3) if n is not divisible by number between 2, sqrt(n)
        if n <= 1:
            return False
        if n <= 3:
            return True
        limit = int(math.sqrt(n)) + 1
        for d in range(2, limit):
            if n % d == 0:
                return False
        return True
    for num in itertools.count(start=2):
        if is_prime(num):
            yield num
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

<!-- grade_id: ruler, points: 0, grade: false, solution: true -->
```python
## %%pycodestyle
def ruler():
    # YOUR CODE HERE
    start = 1
    count = 1
    while True:
        for num in range(start, start + count):
            yield from itertools.repeat(num, count)
        start = start + count
        count = count + 1
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

```python
## %%pycodestyle
def look_say():
    # YOUR CODE HERE
    entry = "1"
    while True:
        yield entry
        next_entry = ''
        count = 1
        
        for i in range(1, len(entry)):
            if entry[i] == entry[i - 1]:
                count += 1
            else:
                next_entry += str(count) + entry[i - 1]
                count = 1
        next_entry += str(count) + entry[-1]
        entry = next_entry
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
## %%pycodestyle
def enc_caesar(s: str, k: int) -> str:
    '''
    Encrypt the message s using Ceasar cipher with key k.
    
    >>> enc_caeasar('thequickbrownfoxjumpsoverthelazydog', 1)
    'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'
    '''
    # YOUR CODE HERE
    encrypted = []
    k = k % 26
    for char in s:
        if 'a' <= char <= 'z':
            shifted = chr((ord(char) - ord('a') + k) % 26 + ord('a'))
            encrypted.append(shifted)
        else:
            encrypted.append(char)
    return ''.join(encrypted)
```

Output:
```
5:1: W293 blank line contains whitespace
19:1: W391 blank line at end of file
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
## %%pycodestyle
def dec_caesar(s: str, k: int) -> str:
    '''
    Decrypt the message s using Ceasar cypher with key k.
    
    >>> dec_caeasar('uifrvjdlcspxogpykvnqtpwfsuifmbazeph', 1)
    'thequickbrownfoxjumpsoverthelazydog'
    '''
    # YOUR CODE HERE
    decrypted = []
    k = k % 26
    for char in s:
        if 'a' <= char <= 'z':
            shifted = chr((ord(char) - ord('a') - k) % 26 + ord('a'))
            decrypted.append(shifted)
        else:
            decrypted.append(char)
    return ''.join(decrypted)
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.


<!-- grade_id: is_valid, points: 0, grade: false, solution: true -->
```python
## %%pycodestyle
def is_valid(s):
    # YOUR CODE HERE
    stack = []
    matching_bracket = {'(': ')', '{': '}', '[': ']'}

    for char in s:
        if char in matching_bracket:
            stack.append(char)
        elif char in matching_bracket.values():
            if not stack or matching_bracket[stack.pop()] != char:
                return False
        else:
            return False
    return len(stack) == 0
```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```

```python

```